---
  title: "Test"
  description: "Test"
---

## An exercise title written in sentence case

```yaml
type: NormalExercise 
lang: python
xp: 100 
skills: 2
key: c67c169ecc   
```


One of the key skills within Data Science is being able to visualize data in order to better understand it. One of the main packages for plotting within Python is called Matplotlib.

Before we use any of the methods in python, we need to import the package.
Most coders abbreviate this package to plt and therefore it is good practice to do the same in order for other people to understand your code.
To import Matplotlib we write the following line in python:
import matplotlib.pyplot as plt


`@instructions`
- Instruction 1

- Instruction 2
- Instruction 3

`@hint`
- Here is the hint for this setup problem. 
- It should get students 50% of the way to the correct answer.
- So don't provide the answer, but don't just reiterate the instructions.
- Typically one hint per instruction is a sensible amount.

`@pre_exercise_code`

```{python}
# Load datasets and packages here.
```

`@sample_code`

```{python}
# Your
# sample
# code
# should
# be
# ideally
# 10 lines or less,
# with a max
# of 16 lines.
```

`@solution`

```{python}
# Answer goes here
# Make sure to match the comments with your sample code
# to help students see the differences from solution
# to given.
```

`@sct`

```{python}
# Update this to something more informative.
success_msg("Some praise! Then reinforce a learning objective from the exercise.")
```
